---
evaluation_interval: 1m

rule_files:
  - ../rules/smartmon_disk_health.yaml

tests:
  - interval: 1m
    input_series:
      - series: obmondo_monitoring{certname="jacen.enableit", alert_id="monitor::system::service::smartmon::disk_healthy"}
        values: 1x1000
      - series: smartmon_device_smart_healthy{certname="jacen.enableit", disk="/dev/sdb"}
        values: 0x1000

    alert_rule_test:
      - alertname: 'monitor::system::service::smartmon::disk_healthy'
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: 'critical'
              certname: 'jacen.enableit'
              disk: '/dev/sdb'
              alert_id: 'monitor::system::service::smartmon::disk_healthy'
            exp_annotations:
              summary: 'Disk **/dev/sdb** on **jacen.enableit** has SMART failure'
              description: 'Disk **/dev/sdb** on **jacen.enableit** has SMART failure'
